Quick-quick-slow: the foxtrot migration and dynamic non-breeding range of the Arctic raptor

Author response:

The following is the authors’ response to the current reviews.

Reviewer #4

We sincerely appreciate the time and effort you have taken to review our manuscript. We followed your recommendations to polish the text and make it easier to understand.

Regarding terms and terminology, we changed “non-breeding” everywhere in the text to “over- wintering.”

Regarding the title, as it was suggested by reviewer #1 as his recommendation, we tried to find a compromise and make the changes you suggested but left part of the suggestion from reviewer #1. So, now it’s “Foxtrot migration and dynamic over-wintering range of an arctic raptor”

Thank you for highlighting the importance of snow cover and changes in snow cover as a possible factor of over-wintering movements. We appreciate your feedback and have explored several approaches to address this issue. Specifically, we examined how both snow cover extent and changes in snow cover influenced movement distance. However, we found no effect of either factor on movement distance.

Our data show that birds leave their sites in October and move southwest, even though snow cover is minimal at that time. They also leave their sites in November and in subsequent months, regardless of the snow cover levels. Thus, we observed no pattern of birds leaving sites when snow cover reaches a specific threshold (e.g., 75-80%). Similarly, we found no evidence of birds staying in areas with a certain snow cover extent (e.g., 30%), nor did they leave sites when snow cover increased by a specific amount (e.g., by 10 or 20%).

It is possible that more experienced birds anticipate that October plots will become inaccessible later in the winter and, therefore, leave early without waiting for significant snow accumulation. Alternatively, other factors, such as brief heavy snowfalls, may trigger movement, even if these do not lead to sustained increases in snow cover. Multiple factors, possibly acting asynchronously, could also play a role. This complexity adds an interesting dimension to the study of ecological patterns. However, in this study, we chose to focus on describing the migration pattern itself and its impact on aspects like over-winter range determination and population dynamics. While we have prioritized this approach, we remain committed to further analyzing the data to uncover additional details about this behavior.

In response to your suggestion, we have expanded the Methods sections to clarify that we tested the effects of snow cover and changes in snow cover on distance (Lines 241-246); the Results section (Lines 348-349). We have also included the relevant plots in the Supplementary Materials. In the Discussion, we noted that this approach did not reveal any significant dependence and acknowledged that this issue requires further investigation (Lines 422-459).

---------

The following is the authors’ response to the previous reviews.

Reviewer #2:

We sincerely appreciate the time and effort you have taken to review our manuscript.

First of all, we apologize for publishing the preprint without incorporating certain adjustments outlined in our earlier response, particularly in the Methods section. This was due to an oversight regarding the different versions of the manuscript. We have corrected this mistake. Our response to the feedback on this section (Methods), with line numbers of the changes made, is immediately below this response. In addition, we have included the units of measurement (mean and standard deviation) in both the results and figure captions for clarity.

To focus on the main point regarding wintering strategies, we acknowledge that in the previous versions, this aspect was inadequately addressed and caused some confusion. In the revised edition, both the Introduction and the Discussion have been thoroughly reworked.

As you suggested, we have removed the long introductory paragraph and all references to foxtrot migrations from the Introduction. As a result, the Introduction is now short and to the point. In the second paragraph, we explain why we propose the wintering strategies outlined (L74-81).

In the Discussion, we've added a substantial new section at the beginning that discusses different wintering strategies. We have also updated Figure 4 accordingly. Previously, we erroneously suggested that Montagu's harrier and other African-Palaearctic migrants might adopt wintering strategies similar to those we describe. Upon further investigation, however, we found that almost all African-Palaearctic migrants exhibit an itinerant wintering strategy. Conversely, the strategy we describe is primarily observed in mid-latitude wintering species.

We have shown that, unlike itinerancy, the birds in our study don't pause for 1-2 months at multiple non-breeding sites, but instead migrate significant distances, up to 1000 km, throughout the winter. Furthermore, unlike itinerancy, the sites they reach are consistently snow-free throughout the year. Following the logic of publications on Montagu's harriers (Schlaich et al. 2023), our birds do not wait for favorable conditions at the next site, as is typical of itinerancy. Moreover, this behavior is influenced by external factors such as snow cover dynamics and occurs primarily in mid-latitudes. Researchers studying a species similar to our subject, the Common buzzard, observed a similar pattern and termed it "prolonged autumn migration" rather than itinerancy. Although their transmitters stopped working in mid-winter, precluding a full observation of the annual cycle, they captured the essence of continued migration at a slower pace, distinct from itinerancy. We've detailed all of these findings in a new section.

In addition, we acknowledge the mischaracterization of the implications of our research as ‘Conservation implications’ and have corrected this to ‘Mapping ranges and assessing population trends’, as you suggested.

Finally, we've rewritten the Conclusion, removing overly grandiose statements and simply summarizing the main findings.

We appreciate your time and effort in reviewing our manuscript. With your invaluable input, it has become clearer, more concise, and easier to understand.

Dataset: unclear what is the frequency of GPS transmissions. Furthermore, information on relative tag mass for the tracked individuals should be reported.

We have included this information in our manuscript (L 115-122). We also refer to the study in which this dataset was first used and described in detail (L 123).

Data pre-processing: more details are needed here. What data have been removed if the bird died? The entire track of the individual? Only the data classified in the last section of the track? The section also reports on an 'iterative procedure' for annotating tracks, which is only vaguely described. A piecewise regression is mentioned, but no details are provided, not even on what is the dependent variable (I assume it should be latitude?).

Regarding the deaths, we only removed the data when the bird was already dead. We estimated the date of death and excluded tracking data corresponding to the period after the bird's death. We have corrected the text to make this clear (L 130-131).

Regarding the piecewise regression. We have added a detailed description on lines 136-148.

Data analysis: several potential issues here:

(1) Unclear why sex was not included in all mixed models. I think it should be included.

Our dataset contains 35 females and eight males (L116). This ratio does not allow us to include sex in all models and adequately assess the influence of this factor. At the same time, because adult females disperse farther than males in some raptor species, we conducted a separate analysis of the dependence of migration distance on sex (Table S8) and found no evidence for this in our species. We have written about that in the Methods (L177-181) and after in the Results (L277-278).

(2) Unclear what is the rationale of describing habitat use during migration; is it only to show that it is a largely unsuitable habitat for the species? But is a formal analysis required then? Wouldn't be enough to simply describe this?

Habitat use and snow cover determine the two main phases (quick and slow) of the pattern we describe. We believe that habitat analysis is appropriate in this case, and a simple description would be uninformative and not support our conclusions.

(3) Analysis of snow cover: such a 'what if' analysis is fine but it seems to be a rather indirect assessment of the effect of snow cover on movement patterns. Can a more direct test be envisaged relating e.g. daily movement patterns to concomitant snow cover? This should be rather straightforward. The effectiveness of this method rests on among-year differences in snow cover and timing of snowfall. A further possibility would be to demonstrate habitat selection within the entire non-breeding home range of an individual in relation snow cover. Such an analysis would imply associating presenceabsence of snow to every location within the non-breeding range and testing whether the proportion of locations with snow is lower than the proportion of snow of random locations within the entire nonbreeding home range (95% KDE) for every individual (e.g. by setting a 1/10 ratio presence to random locations).

The proposed analysis will provide an opportunity to assess whether the Rough-legged buzzard selects areas with the lowest snow cover, but will not provide an opportunity to follow the dynamics and will therefore give a misleading overall picture. This is especially true in the spring months. In March-April, Rough-legged buzzards move northeast and are in an area that is not the most open to snow. At this time, areas to the southwest are more open to snow (this can be seen in Figure 3b). If we perform the proposed analysis, the control points for this period would be both to the north (where there is more snow) and to the south (where there is less snow) from the real locations, and the result would be that there is no difference in snow cover.

A step-selection analysis could be used, as we did in our previous work (Curk et al 2020 Sci Rep) with the same Rough-legged buzzards (but during migration, not winter). But this would only give us a qualitative idea, not a quantitative one - that Rough-legged Buzzards move from snow (in the fall) and follow snowmelt progression (in the spring).

At the same time, our analysis gives a complete picture of snow cover dynamics in different parts of the non-breeding range. This allows us to see that if Rough-legged buzzards remained at their fall migration endpoint without moving southwest, they would encounter 14.4% more snow cover (99.5% vs. 85.1%). Although this difference may seem small (14.4%), it holds significance for rodent-hunting birds, distinguishing between complete and patchy snow cover.

Simultaneously, if Rough-legged buzzards immediately flew to the southwest and stayed there throughout winter, they would experience 25.7% less snow cover (57.3% vs. 31.6%). Despite a greater difference than in the first case, it doesn't compel them to adopt this strategy, as it represents the difference between various degrees of landscape openness from snow cover.

Changes from the earlier version to the revised version

Abstract

Abstract

Added in the revised version: Advances in tracking technologies have revealed the diverse migration patterns of birds, which are critical for range mapping and population estimation. While population trends are typically estimated in breeding ranges where birds remain stationary, trends are commonly assessed in non-breeding ranges for species that breed in inaccessible regions such as the Arctic. This task is challenging due to the extensive movements of birds in these ranges, which require a deep understanding of the dynamics of these movements. However, these movements remain understudied, particularly in the mid-latitudes, where many Arctic breeders winter, increasing the risk of misestimating their ranges and numbers. Here, we show that the Arctic breeding raptor Rough-legged buzzard, which overwinters in the mid-latitudes, has a specific wintering strategy: during the entire non-breeding season, they move about 1000 km back and forth under the influence of the advancing snowline, making their non-breeding area dynamic. In essence, this movement represents an extension of the quick migration process, albeit at a slower pace, and we have termed this migration pattern ‘foxtrot migration,’ drawing an analogy to the alternating fast and slow movements of the foxtrot dance. These results highlight the potential errors in range mapping from single mid-winter surveys and emphasize the importance of this migration pattern in assessing the conservation status of bird species. Understanding this migration pattern helps to correctly estimate bird populations in non-breeding areas, which is especially important for species that nest in hard-to-reach regions such as the Arctic.

Removed in the revised version: Advancements in tracking technologies have revolutionized our understanding of bird migrations, revealing a diverse array of migratory behaviors. We propose a new pattern of migratory behavior termed ‘foxtrot migration,’ characterized by alternating quick and slow phases during the non-breeding period. This behavior involves directional and continuous movements, distinct from traditional seasonal itinerancy or forage-on-fly migration. Using the Rough-legged buzzard as a model species, we confirmed the presence of foxtrot migration and ‘dynamic non-breeding ranges’, driven by environmental factors such as snow cover dynamics. We advocate for accurate representation of dynamic ranges on maps and emphasize the need to consider range dynamics when assessing species conservation status. Our findings underscore the importance of understanding complex migratory behaviors in the face of environmental change, facilitated by advancements in tracking technologies. This knowledge is crucial for effective conservation strategies amid ongoing global environmental challenges.

Introduction

Added in the revised version: In recent years, advances in tracking technology have greatly improved our understanding of bird migration (Wikelski et al. 2007, Kays et al. 2015, Jetz et al. 2022, Kays and Wikelski 2023). By tracking tagged birds, we can study various migratory behaviors, such as their propensity to migrate, connectivity and philopatry, length of migration routes, variations in flight speed, and movements within a single season (Alerstam et al. 2003, Newton 2008, Alerstam 2011, Berthold et al. 2013, Trierweiler et al. 2013, Schlaich et al. 2023). Different combinations of these behaviors contribute to various bird migration patterns (Chapman et al. 2014, Alerstam and Bäckman 2018, Lislevand et al. 2020). Understanding the migratory patterns of birds is critical to accurately mapping their ranges and estimating population trends. Typically, population trends are assessed in breeding areas, where birds are linked to their nesting sites, allowing reliable estimates of breeding numbers (Reif 2013). However, access to some breeding areas, such as the Arctic, presents logistical and political challenges (Gallo-Cajiao et al. 2023, Koivurova and Shibata 2023). As a result, population trends for these species are often estimated in non-breeding areas where birds aren’t tied to specific locations (Robinson et al. 2005, JIa et al. 2016). This challenges accurately censusing birds and delineating their ranges, especially when birds move extensively in non-breeding ranges during the season, and censuses are conducted only once per non-breeding season. Therefore, understanding bird movements in non-breeding areas is essential for understanding the ecology of birds in these regions and has practical implications for accurately assessing their ranges and trends in their populations.

Added in the revised version: Many bird species experience significant seasonal changes in food availability within their non-breeding range due to environmental conditions. This often triggers movements across these areas. For Palearctic migrants wintering in Africa, these movements have been studied extensively (Moreau 1972, Trierweiler et al. 2013, Thorup et al. 2017, Schlaich et al. 2023). However, the non-breeding movements of mid-latitude birds remain poorly understood. At the same time, many Arctic birds winter in the mid-latitudes, and understanding their movement patterns could be crucial for interpreting their population trends. In the mid-latitudes, food availability is strongly influenced by environmental factors such as snow cover dynamics. The progressive movement of the snow cover line from north (-east) to south (-west) and back again between October and May in many mid-latitude regions, particularly those exposed to northwesterly wind systems, significantly affects food availability for numerous bird species (Sonerud 1986, Vansteelant et al. 2011). We propose that species dependent on this environmental factor gradually move away from snow-covered areas during winter and then gradually move back in the opposite direction, resulting in a directed and continuous displacement of their non-breeding range.

Removed in the revised version: In recent decades, advanced tracking technology has greatly improved our knowledge of bird migrations (Wikelski et al. 2007; Jetz et al. 2022; Kays & Wikelski 2023). In the 19th century, the finding of a White stork (Ciconia ciconia) in Germany with an African arrow led to the discovery that birds did not hibernate at the bottom of water bodies during the winter but rather made extensive seasonal migrations (Richter & Bick 2018). Subsequent studies of bird migration in the 20th century using ringing techniques revealed a diversity of seasonal migratory behaviors, including 1) propensity to migrate (migrants, residents, partial migrants), 2) differences among populations (leapfrog, chain, parallel migrations), 3) connectivity and philopatry (constancy of breeding and non-breeding locations), 4) length of migration routes (long-distance and short-distance migrants), 5) variations in flight paces during migration (with stopovers, forage-on-fly, nonstop), and 6) tendencies to move within a single season (breeding and non-breeding itinerancy) (Alerstam, Hedenström & Åkesson 2003; Newton 2008; Alerstam 2011; Berthold, Gwinner & Sonnenschein 2013). Different combinations of these behaviors contribute to a wide variety of bird migration patterns (Chapman et al. 2014; Alerstam & Bäckman 2018; Lislevand et al. 2020). However, the advent of 21st-century technology, such as GPS transmitters, provides a higher resolution of movement and the ability to observe more details of migratory behavior (Wikelski et al. 2007; Jetz et al. 2022). Just as advances in image resolution allow us to see finer details in images, improvements in bird tracking technology now allow us to identify previously undiscovered migratory behaviors, revealing new patterns.

Removed in the revised version: Current technologies allow observing forms of migratory behavior that are intermediate or integrative to those described above. Itinerancy was initially defined as the movements of Afro-Palearctic migrants during the non-breeding season in the Sahel, with a combination of stopovers and flights in between (Moreau 1972). Subsequently, researchers investigated this behavior for multiple species wintering in Africa, attributing their movements during the non-breeding season to changes in food availability predicted by NDVI (Zwarts et al. 2012; Schlaich et al. 2016; Thorup et al. 2017; Verhoeven et al. 2021). In 1998, Garcia and Arroyo hypothesized that certain species, like the Montagu’s harrier (Circus pygargus), might not only cover tens to hundreds of kilometers between various sites in their winter range but could also demonstrate directional, continuous movement throughout the entire non-breeding season. This movement may span a distance similar to their migration distance (ca. 1000 km), essentially indicating a prolonged, slower-paced continuation of migration (García & Arroyo 1998). Because this behavior is intermediate between seasonal movement (itinerancy) and the actual migration process, we can describe it as either ‘persistent directional non-breeding itinerancy’ or ‘very slow forage-on-fly migration in the non-breeding area’. Both options are clumsy and can confuse. We suggest naming this behavior with a new term to avoid confusion with established terms such as ‘itinerancy’ for any seasonal movement and ‘forage-on-fly’ denoting movement from breeding to non-breeding sites (Alerstam, Hake & Kjellén 2006; Alerstam 2011; Lopez-Ricaurte et al. 2024). Considering this behavior as an extension of migration and looking at the complete life cycle of such species, it includes a quick phase during the transition between breeding and non-breeding areas, followed by a slow phase of directed and seasonal movement within the non-breeding area, and finally another quick phase towards the breeding area. Analogous to the alternating fast and slow movements of the foxtrot dance, we propose the term ‘foxtrot migration’ for this seasonal movement to provide a concise and easily understandable description (Figure 1a). Consequently, we suggest to term the non-breeding range of species exhibiting this migration type as ‘dynamic range’ (Figure 1b).

Removed in the revised version: Foxtrot migration and dynamic winter range scheme.

Removed in the revised version: a) Foxtrot migration. The color bar along the y-axis corresponds to the color coding of the habitats in panel b of the chart on the right. b) Dynamic winter range.

Removed in the revised version: The concept of dynamic range, derived from foxtrot migration, is critical to understanding the current range declines and shifts observed in many species. Typically, a species’ range is divided into areas of common and uncommon, with temporal heterogeneity limited to breeding and non-breeding distinctions (Keller et al. 2020; Billerman et al. 2024). However, suppose a species exhibits a foxtrot migration, moving gradually in one direction and back again during the non-breeding season. In this case, it becomes inaccurate to represent its non-breeding range as static. This discrepancy is particularly pronounced when ranges are mapped using mid-winter surveys, such as Christmas counts. During this period, the species may be present at the outer limits of its non-breeding range, leading to a false classification as common in that area when, in fact, it is rare if the entire non-breeding season is considered. Similarly, extended negative population dynamics of a species in a particular part of its non-breeding range could be misinterpreted as a threat to the entire population. Indeed, if this species has a foxtrot migration, range dynamics may have changed, possibly resulting in the species ceasing to migrate to that particular part of its range. Advanced technologies now make it possible the identification of foxtrot migration and range dynamics. This, in turn, facilitates accurate range mapping and proper assessment of threats to the species.

Removed in the revised version: This pattern of behavior has been proposed in species that overwinter in Africa, where non-breeding movements are primarily related to the distribution of insect abundance (García & Arroyo 1998). At the same time, it is reasonable to assume that this behavior may also be characteristic of species that spend the non-breeding season in mid-latitudes, where food availability may be linked to snow cover dynamics. These dynamics have changed particularly rapidly in recent decades due to global warming (Davy & Outten 2020; Previdi, Smith & Polvani 2021). As a result, global climate change may significantly affect birds that breed in the Arctic and winter in mid-latitudes. The Arctic is experiencing more pronounced effects of global change than other regions, with declines in the abundance and range of several species and populations (Post et al. 2009; Gilg et al. 2012; Schmidt et al. 2012). Therefore, understanding how snow cover dynamics affect Arctic species across habitats throughout the annual cycle is essential today. Furthermore, elucidating the foxtrot migration and dynamic non-breeding range of mid-latitude wintering Arctic birds may provide valuable information on expected future changes in their non-breeding range.

Earlier version: We used the Rough-legged buzzard (Buteo lagopus) as a model species to investigate this phenomenon. The Rough-legged buzzard is an Arctic breeding and mid-latitude wintering raptor (Ferguson-Lees & Christie 2001; Bechard & Swem 2002). Rough-legged buzzards feed mainly on small rodents in the Arctic and its wintering grounds (Tast, Kaikusalo & Lagerström 2010; Pokrovsky et al. 2014). They prefer open areas for hunting, and trees and tall bushes or uplands for resting. In the Arctic, such areas are the southern and typical tundra (Walker et al. 2005), and in the mid-latitudes, areas with fields and patches of forests (wooded fields). The taiga zone, where there is little open space, is unsuitable for them, although they may nest in the northern taiga zone on the border with the tundra zone (Sundell et al. 2004). Snow cover and day length play an important role in their life (Terraube et al. 2015; Curk et al. 2020; Pokrovsky et al. 2021). Rough-legged buzzards can only hunt during the day (Pokrovsky et al. 2021), and heavy snow cover makes hunting for small rodents problematic (Sonerud 1986; Vansteelant, Faveyts & Buckens 2011). These two environmental factors, therefore, affect the availability of prey for Rough-legged buzzards. At the same time, these factors vary considerably in mid-latitudes during winter. Thus, prey availability for Rough-legged buzzards in the mid-latitudes increases until mid-winter if they migrate southwards and after mid-winter if they migrate northwards. We, therefore, assume that Rough-legged buzzards could track prey availability and experience a foxtrot migration and dynamic non-breeding range.

Revised version: The Rough-legged buzzard (Buteo lagopus) is an Arctic breeding and mid-latitude wintering raptor (Ferguson-Lees and Christie 2001, Bechard and Swem 2002). Rough-legged buzzards feed mainly on small rodents in the Arctic and its wintering grounds (Tast et al. 2010, Pokrovsky et al. 2014). They prefer open areas for hunting, and trees and tall bushes or uplands for resting. In the Arctic, such areas are the southern and typical tundra (Walker et al. 2005), and in the mid-latitudes, areas with fields and patches of forests (wooded fields). The taiga zone, where there is little open space, is unsuitable for them, although they may nest in the northern taiga zone on the border with the tundra zone (Sundell et al. 2004). Snow cover and day length are important in their life (Terraube et al. 2015, Curk et al. 2020, Pokrovsky et al. 2021). Rough-legged buzzards can only hunt during the day (Pokrovsky et al. 2021), and heavy snow cover makes hunting for small rodents problematic (Sonerud 1986, Vansteelant et al. 2011). These two environmental factors, therefore, affect the availability of prey for Rough-legged buzzards. At the same time, these factors vary considerably in mid-latitudes during winter. Thus, prey availability for Rough-legged buzzards in the mid-latitudes increases until mid-winter if they migrate southwards and after mid-winter if they migrate northwards. We, therefore, assume that Rough-legged buzzards could track prey availability and experience a dynamic non-breeding range.

Earlier version: For this study, we hypothesized that Rough-legged buzzards undertake a combined foxtrot migration and dynamic non-breeding range in response to seasonal changes in mid-latitude environmental factors. We made the following predictions. 1) Buzzards would exhibit a directional and seasonal movement pattern during the non-breeding period, moving from the northeast to the southwest and back again. This would result in a dynamic non-breeding range that would continue to move geographically throughout the season. 2) Non-breeding movements would occur in suitable open habitats, whereas fall and spring migrations would occur in unfitting forested areas. 3) Non-breeding movements would differ from fall and spring migrations in duration, extent, speed, and direction, however would continue throughout the season. 4) During non-breeding movements, Rough-legged buzzards will experience less snow cover than if they had stayed where they arrived at the end of the fall migration.

Revised version: In this study, we tracked GPS-tagged Rough-legged buzzards for ten years and analyzed their movements in the non-breeding mid-latitude range in relation to snow cover. For this study, we made the following predictions. 1) Rough-legged buzzards would exhibit a directional and seasonal movement pattern during the non-breeding period, moving from the northeast to the southwest and back again. This would result in a dynamic non-breeding range that would continue to move geographically throughout the season. 2) The non-breeding movements would differ from the fall and spring migrations in duration, extent, speed, and direction but would continue throughout the season. Thus, over the entire annual cycle, Rough-legged buzzards would be characterized by a migratory pattern consisting of an alternation of quick (fall and spring migrations) and slow (non-breeding movements) phases. 3) Non-breeding movements would occur in suitable open habitats, while fall and spring migrations would occur in unsuitable forested areas. Thus, vegetation land cover will determine the extent of quick and slow migration phases. 4) During non-breeding movements, Rough-legged buzzards will experience less snow cover than if they had stayed where they arrived at the end of the fall migration. Thus, snow cover dynamics will determine the dynamics of non-breeding movements.

Removed in the revised version: In the following, we will refer to the spring and fall migrations as the quick phase, the non-breeding movement to the lowest point of latitude as the 1st part of the slow phase, and the movement from the lowest point of latitude to the starting point of the spring migration as the 2nd part of the slow phase (Figure 1).

Material and methods

Dataset

Added in the revised version: For this study, we tracked 43 adult Rough-legged buzzards (35 females and eight males) with the solar GPS-GSM loggers (e-obs GmbH and UKn – University of Konstanz). Twenty-eight birds were fitted with E-obs loggers, 13 birds were fitted with UKn loggers, and two birds were initially fitted with UKn loggers that were later replaced with E-obs loggers. E-obs loggers weighed 45 g, or 3.6% of the bird’s weight, while UKn loggers weighed 15 g, or 1.2% of the bird’s weight. E-obs loggers recorded GPS positions every hour in full battery mode and every 5 hours in normal battery mode. UKn loggers recorded GPS positions every hour with a full battery and every 12 hours with a normal battery. Both models operated continuously 24 hours a day. More detailed information about this dataset can be found in Curk et al. (2022).

Earlier version: For this study, we tracked 43 adult Rough-legged buzzards (35 females and eight males) with the solar GPS-GSM loggers (e-obs GmbH and UKn – University of Konstanz). The fieldwork was carried out in the Russian Arctic in 2013-2019 at four study sites: Kolguev Island (69°16′N, 48°87′E), Nenetsky Nature Reserve (68°20′N, 53°18′E), Vaigach Island (69°43′N, 60°08′E), and Yamal Peninsula (68°12′N, 68°59′E). For details on capture methods and tag parameters, see Curk et al. (2022); for detailed study descriptions, see Pokrovsky et al. (2015) for Kolguev and Pokrovsky et al. (2019) for Yamal and Nenetsky.

Revised version: The fieldwork was carried out in the Russian Arctic in 2013-2019 at four study sites: Kolguev Island (69°16′N, 48°87′E), Nenetsky Nature Reserve (68°20′N, 53°18′E), Vaigach Island (69°43′N, 60°08′E), and Yamal Peninsula (68°12′N, 68°59′E). For details on capture methods, see Curk et al. (2022); for detailed study descriptions, see Pokrovsky et al. (2015) for Kolguev and Pokrovsky et al. (2019) for Yamal and Nenetsky.

Earlier version: During data pre-processing, we estimated the date of death using the accelerometer or GPS data and removed the tracking data if the bird was dead. We then removed duplicated timestamps and calculated the mean daily positions of each individual. We partitioned the resulting dataset into several periods: 1) breeding, 2) fall migration, 3)1st part of winter, 4) 2nd part of winter, and 5) spring migration. We estimated the migration dates – the start and stop dates of the spring and fall migrations – using an iterative search procedure for piecewise regression described by Crawley (2007). We estimated the date between winter’s first and second parts as the day when the mean daily latitude was minimum.

Revised version: During data pre-processing, we estimated the date of death using an accelerometer and GPS data and excluded tracking data corresponding to the period after the bird’s death. We then removed duplicated timestamps and calculated the mean daily positions of each individual. We partitioned the resulting dataset into several periods: 1) breeding, 2) fall migration, 3)1st part of winter, 4) 2nd part of winter, and 5) spring migration. We estimated the migration dates – the start and stop dates of the spring and fall migrations – using an iterative search procedure for piecewise regression described by Crawley (2007). We created two models for each of the four dates: a simple model and a piecewise one. The simple model was linear, with latitude as the response variable and day of the year as a fixed effect. The piecewise model was also linear, with latitude as the response variable and day of the year as a fixed effect, and included two logical statements in its formula because we expected two linear segments in the fit. Specifically, it used “Day Of the Year < BREAK” to define the left regression (before the estimated date) and “Day Of the Year ≥ BREAK” for the right regression (after the estimated date). We explored different BREAK values to fit the piecewise model, looking for the value that minimized the residual standard error. The BREAK values examined included those associated with the start and end of the fall migration (September 5 to October 5 and October 5 to November 5, respectively), as well as those associated with the start and end of the spring migration (April 10 to May 10 and May 10 to June 10, respectively). We then validated (using the R function ‘anova’) that the piecewise model significantly improved the fit compared to the simple model (p<0.001). We estimated the date between winter’s first and second parts as the day when the mean daily latitude was minimum.

Data analysis

First, we used linear mixed-effects models (R function ‘lmer’ in the library ‘lme4’ (Bates et al. 2015)) to investigate whether or not Rough-legged buzzards migrated during winter. Latitude was the response variable, day of the year was a fixed effect, and individuals and year were included as random effects. Analyses were conducted separately for each migration period (fall, first phase of winter, second phase of winter, and spring). For both phases of the winter migration, we analyzed two additional models with longitude as the response variable instead of latitude. Likelihood ratio tests were used to compare candidate models. The year was not a calendar year but a year between two consecutive breeding seasons. Thus, fall migration, consecutive winter, and consecutive spring have the same value for the year. The day of the year was recalculated consecutively.

Earlier version: Second, we used linear mixed-effects models (R function ‘lmer’ in the library ‘lme4’) to investigate whether migrations’ parameters differ between the migration periods. We analyzed four migration parameters: distance, duration, speed, and direction. The distance was calculated as the distance between two coordinates (start and end of migration) using the R function ‘distm’ in the library ‘geosphere’ (Hijmans 2016). The duration was calculated as the number of days between the start and end of migration. Speed was calculated as the ratio of distance to duration. The direction was calculated as the bearing from the start of the migration coordinates to the end of the migration coordinates using the R function ‘bearing’ in the library ‘geosphere’ (Hijmans 2016). The migration parameter was used as the response variable, the type of migration as a fixed factor, and individuals as a random factor. Likelihood ratio tests were used to compare candidate models. We considered four different parameters of migration (distance, duration, speed, and direction) and four types of migration (fall, first phase of winter, second phase of winter, and spring). The analysis was done separately for each of the migration parameters. Then, we used post hoc comparisons using the R function ‘emmeans’ in the library ‘emmeans’ (Lenth et al. 2019) to compare the estimated means. In some raptor species, adult females disperse further than males (Mearns & Newton 1984; Serrano et al. 2001; Bildstein 2006; Whitfield et al. 2009). Therefore, we conducted an additional analysis on the effect of sex on migration length using linear mixed-effects models (R function ‘lmer’ in the library ‘lme4’). The migration distance was used as the response variable, sex as a fixed factor, and individuals as a random factor.

Revised version: Second, we used linear mixed-effects models (R function ‘lmer’ in the library ‘lme4’) to investigate whether migrations’ parameters differ between the migration periods. We analyzed four migration parameters: distance, duration, speed, and direction. The distance was calculated as the distance between two coordinates (start and end of migration) using the R function ‘distm’ in the library ‘geosphere’ (Hijmans 2016). The duration was calculated as the number of days between the start and end of migration. Speed was calculated as the ratio of distance to duration. The direction was calculated as the bearing from the start of the migration coordinates to the end of the migration coordinates using the R function ‘bearing’ in the library ‘geosphere’ (Hijmans 2016). The migration parameter was used as the response variable, the type of migration as a fixed factor, and individuals as a random factor. Likelihood ratio tests were used to compare candidate models. We considered four different parameters of migration (distance, duration, speed, and direction) and four types of migration (fall, first phase of winter, second phase of winter, and spring). The analysis was done separately for each of the migration parameters. Then, we used post hoc comparisons using the R function ‘emmeans’ in the library ‘emmeans’ (Lenth et al. 2019) to compare the estimated means.

Added in the revised version: In some raptor species, adult females disperse further than males (Mearns and Newton 1984, Serrano et al. 2001, Bildstein 2006, Whitfield et al. 2009). Therefore, we conducted an additional analysis on the effect of sex on migration length using linear mixed-effects models (R function ‘lmer’ in the library ‘lme4’). The migration distance was used as the response variable, sex as a fixed factor, and individuals as a random factor.

Earlier version: Third, we investigated whether vegetation land cover differed between areas crossed during the quick (fall and spring) and slow (winter) migrations. We used the combined Terra and Aqua Moderate Resolution Imaging Spectroradiometer (MODIS) Land Cover Climate Modelling Grid (CMG) (MCD12C1) version 6 dataset (Friedl & Sulla-Menashe 2015). We used a modified Leaf Area Index (LAI) as a classification scheme. We combined all four forest types and savannas into one category (forest) and excluded the categories: water bodies, and unclassified. We, therefore, had five types of vegetation cover: forest, grassland, cropland, shrubland, and urban. We annotated the mean daily positions with the vegetation cover type using the Env-DATA tool (Dodge et al. 2013). We used general linear mixed effects models with a binomial distribution (R function ‘glmer’ in the library ‘lme4’) to investigate whether vegetation cover types differ between migration periods. Presence/absence of the studied vegetation land cover type was used as a response variable, migration type as a fixed factor, and individuals as a random factor. The analysis was done separately for each of the vegetation land cover types.

Revised version: Third, we investigated whether vegetation land cover differed between areas crossed during the quick (fall and spring) and slow (winter) migrations. We used the combined Terra and Aqua Moderate Resolution Imaging Spectroradiometer (MODIS) Land Cover Climate Modelling Grid (CMG) (MCD12C1) version 6 dataset (Friedl and Sulla-Menashe 2015). We used a modified Leaf Area Index (LAI) as a classification scheme. We combined all four forest types and savannas into one category (forest) and excluded the categories: water bodies, and unclassified. We, therefore, had five types of vegetation cover: forest, grassland, cropland, shrubland, and urban. We annotated the mean daily positions with the vegetation cover type using the Env-DATA tool (Dodge et al. 2013). We used general linear mixed effects models with a binomial distribution (R function ‘glmer’ in the library ‘lme4’) to investigate whether vegetation cover types differ between migration periods. Presence/absence of the studied vegetation land cover type was used as a response variable, migration type as a fixed factor, and individuals as a random factor. The analysis was done separately for each of the vegetation land cover types.

Fourth, we investigated whether snow cover could drive the slow migration phenomenon. We then compared the snow cover conditions the birds experienced during the winter with two hypothetical snow cover conditions that the birds would have experienced if they had not migrated during the winter. The first hypothetical snow cover condition would have happened if the birds had stayed where they arrived from the north (i.e., where their fall migration ended). To estimate this parameter, we calculated the winter dynamics of the average snow cover at the minimum convex polygons (MCP) occupied by the birds in October and April (northeast of their winter range). A second hypothetical snow cover condition would be if the birds flew immediately to the southwest and spent the whole winter there. To evaluate this, we calculated the winter dynamics of average snow cover on the MCPs occupied by the birds in January and February (southwest of their winter range). We then compared the values obtained for the real snow cover and two hypothetical snow covers using general linear mixed effects models with a binomial distribution (R function ‘glmer’ in the ‘lme4’ library). Presence/absence of snow cover was used as a response variable, type of snow cover (real, 1st hypothetical, or 2nd hypothetical) as a fixed factor, and years as a random factor. The analysis was done separately for each month. We then used post hoc comparisons to compare the estimated means, using the R function ‘emmeans’ in the ‘emmeans’ library (Lenth et al. 2019).

Earlier version: We obtained monthly snow cover data with a spatial resolution of ca 500 meters (Global SnowPack MODIS) from the German Aerospace Center (DLR). This product is based on the Moderate Resolution Imaging Spectroradiometer (MODIS) daily snow cover products MOD10A1 and MYD10A1 (version 6 as provided by the National Snow and Ice Data Center NSIDC), which have been processed to remove the gaps due to cloud cover and polar darkness (Dietz, Kuenzer & Dech 2015). These processing steps include a combination of data available from different satellites (Aqua and Terra), 3-day temporal moving window filtering, a regional snow line elevation interpolation relying on a Digital Elevation Model (DEM), and a seasonal filter running through the time series for the whole hydrological year (1st of September through August 31st). The proportion of days in which one pixel is snow-covered per month is referred to here as fractional snow cover and is derived from these daily gap-filled rasters. Five MODIS tiles (h19v03, h20v03, h20v04, h21v03 and h21v04) were mosaicked and re-projected to WGS84. Then, for each month from October to April, we calculated 95% minimum convex polygons (MCPs) for the distribution of Rough-legged buzzards using the R function ‘mcp’ in library ‘adehabitatHR’ (Calenge 2006). We extracted mean snow cover values from each MCP from every monthly snow cover raster separately, using the R library ‘raster’ (Hijmans & van Etten 2023).

Revised version: We obtained monthly snow cover data with a spatial resolution of ca 500 meters (Global SnowPack MODIS) from the German Aerospace Center (DLR). This product is based on the Moderate Resolution Imaging Spectroradiometer (MODIS) daily snow cover products MOD10A1 and MYD10A1 (version 6 as provided by the National Snow and Ice Data Center NSIDC), which have been processed to remove the gaps due to cloud cover and polar darkness (Dietz et al. 2015). These processing steps include a combination of data available from different satellites (Aqua and Terra), 3-day temporal moving window filtering, a regional snow line elevation interpolation relying on a Digital Elevation Model (DEM), and a seasonal filter running through the time series for the whole hydrological year (1st of September through August 31st). The proportion of days in which one pixel is snow-covered per month is referred to here as fractional snow cover and is derived from these daily gap-filled rasters. Five MODIS tiles (h19v03, h20v03, h20v04, h21v03 and h21v04) were mosaicked and re-projected to WGS84. Then, for each month from October to April, we calculated 95% minimum convex polygons (MCPs) for the distribution of Rough-legged buzzards using the R function ‘mcp’ in library ‘adehabitatHR’ (Calenge 2006). We extracted mean snow cover values from each MCP from every monthly snow cover raster separately, using the R library ‘raster’ (Hijmans and van Etten 2023).

All calculations were performed using R version 4.2.2 ‘Innocent and Trusting’ (R Development Core Team 2022) and RStudio version 353 ‘Elsbeth Geranium’ (Posit team 2022).

Results

Earlier version: Foxtrot migration – always on the move

Revised version: Migration of Rough-legged buzzards – always on the move

Earlier version: Except during the breeding season, Rough-legged buzzard migration continues throughout the year, even after the bird’ arrival at their traditionally recognized ‘wintering ground’ (Figure 2a, b). For both quick and slow phases of the foxtrot migration, linear mixed-effects models with the season as a fixed factor received higher support from the likelihood ratio test (p < 0.001, Tables S1-S3). Rough-legged buzzards started their fall migration (quick phase) on 28 September (hereafter mean±sd for the day of the year: 271±11, n=31) and ended on 12 October (285±11, n=33). The mean latitude/longitude where the birds ended their fall migration was 55.57±1.92°/49.35±5.63° (Figure 2a). During the winter, birds continued to migrate at a slower pace down to 49.53±2.01° latitude (on 5 February, 36±40, n=23) and 34.29±5.11° longitude (on 24 January, 24±47, n=23). Afterward, during the second part of the slow phase, the birds returned to 55.52±2.63° latitude and 49.79±8.24° longitude to start the spring migration (Figure 2a). Rough-legged buzzards started their spring migration to the Arctic on 27 April (117±7, n=27) and arrived at the breeding grounds on 15 May (135±8, n=18).

Revised version: Except during the breeding season, Rough-legged buzzard migration continues throughout the year, even after the birds’ arrival at their traditionally recognized ‘wintering grounds’ (Figure 1a, b). Rough-legged buzzards started their fall migration on 28 September (hereafter mean±sd for the day of the year: 271±11, n=31) and ended on 12 October (285±11, n=33). The mean latitude/longitude where the birds ended their fall migration was 55.57±1.92°/49.35±5.63° (Figure 1a). During the winter, birds continued to migrate at a slower pace down to 49.53±2.01° latitude (on 5 February, 36±40, n=23) and 34.29±5.11° longitude (on 24 January, 24±47, n=23). Afterward, during the second part of the winter, birds returned to 55.52±2.63° latitude and 49.79±8.24° longitude to start the spring migration (Figure 1a). Rough-legged buzzards started their spring migration to the Arctic on 27 April (117±7, n=27) and arrived at the breeding grounds on 15 May (135±8, n=18).

Earlier version: Foxtrot migration of Rough-legged buzzards.

Revised version: Migration of Rough-legged buzzards.

Earlier version: Q – quick phase. Qf – Quick fall phase (orange), Qs – Quick spring phase (yellow), S1 – Slow phase, 1st part (light blue), S2 – Slow phase, 2nd part (dark blue). a) Change in the latitude of 43 Rough-legged buzzards during the year, red line – mean latitude of all birds, black vertical lines – mean dates of start and end of the migration phases, blue vertical line – mean date of the minimum latitude. Grey, sky blue, and piggy pink shaded areas – standard deviation of the means. b) Migration map. c) Difference in the migration parameters between the migration phases. Lines on the direction plot (down, right) represent the mean value for each bird; arrows represent the mean direction.

Revised version: Q – quick phase. Qf – Quick fall phase (orange), Qs – Quick spring phase (yellow), S1 – Slow phase, 1st part (light blue), S2 – Slow phase, 2nd part (dark blue). a) Change in the latitude of 43 Rough-legged buzzards during the year, red line – mean latitude of all birds, black vertical lines – mean dates of start and end of the migration phases, blue vertical line – mean date of the minimum latitude. Grey, sky blue, and piggy pink shaded areas – standard deviation of the means. b) Migration map. c) Difference in the migration parameters between the migration phases. Lines on the direction plot (down, right) represent the mean value for each bird; arrows represent the mean direction for each phase. Boxes on the boxplots show the interquartile range, the whiskers – maximum and minimum values.

Added in the revised version: In the following, we will refer to the spring and fall migrations as the quick phase, the non-breeding movement to the lowest point of latitude as the 1st part of the slow phase, and the movement from the lowest point of latitude to the starting point of the spring migration as the 2nd part of the slow phase (Figure 1a). For both quick and slow phases of the migration, linear mixed-effects models with the season as a fixed factor received higher support from the likelihood ratio test (p < 0.001, Tables S1-S3).

Quick-slow phase features comparison

Earlier version: During quick phase, individual birds flew greater distances in a shorter time, i.e., at a faster rate, than during slow phase. After arriving at what is traditionally known as the wintering grounds, the direction of migration changed, so the direction of quick and slow phases also differed (Table 1, Figure 2c). The quick phase was 1415±50 km long, whereas the slow phase (one part) was 1026±55 km, i.e., 389±60 km shorter (p<0.001, Table S4, Figure 2c). During quick phase, birds flew for 15±3 days, and one part of slow phase lasted 100±4 days, i.e., 85±5 days longer (p<0.001, Table S5, Figure 2c). At the same time, the second part of slow phase was 54±7 days shorter than the first (p<0.001, Table S5, Figure 2c). The migration speed was 104±6 km/day during the quick phase and 12±7 km/day during the slow phase, i.e., about eight times higher (p<0.001, Table S6, Figure 2c). During the fall migration, birds moved in the SSW direction (7±2 deg), then turned 50±3 deg (p<0.001, Table S7, Figure 2c) to the west and started their 1st slow phase until mid-winter. After that, they turned back to the NEE direction (57±2 deg) and performed their 2nd slow phase for several months until they turned 54±3 deg (p<0.001, Table S7, Figure 2c) to the north and started their spring migration. As a result of additional analysis of the effect of sex on migration length, we found no significant difference between the migration distances of males and females (Table S8).

Revised version: During the quick phase, individual birds flew greater distances in a shorter time, i.e., at a faster rate, than during the slow phase. After arriving at what is traditionally known as the wintering grounds, the direction of migration changed, so the direction of quick and slow phases also differed (Table 1, Figure 1c). The quick phase was 1415±50 km long (hereafter mean±sd), whereas the slow phase (one part) was 1026±55 km, i.e., 389±60 km shorter (p<0.001, Table S4, Figure 1c). During the quick phase, birds flew for 15±3 days, and one part of the slow phase lasted 100±4 days, i.e., 85±5 days longer (p<0.001, Table S5, Figure 1c). At the same time, the second part of the slow phase was 54±7 days shorter than the first (p<0.001, Table S5, Figure 1c). The migration speed was 104±6 km/day during the quick phase and 12±7 km/day during the slow phase, i.e., about eight times higher (p<0.001, Table S6, Figure 1c). During the fall migration, birds moved in the SSW direction (7±2 deg), then turned 50±3 deg (p<0.001, Table S7, Figure 1c) to the west and started their 1st slow phase until mid-winter. After that, they turned back to the NEE direction (57±2 deg) and performed their 2nd slow phase for several months until they turned 54±3 deg (p<0.001, Table S7, Figure 1c) to the north and started their spring migration. As a result of additional analysis of the effect of sex on migration length, we found no significant difference between the migration distances of males and females (Table S8).

Vegetation land cover during migration

Earlier version: During quick phases, Rough-legged buzzards cross the forest zone, while during slow phase, they migrate within the grassland and cropland zone (Figure 3). Rough-legged buzzards migrated fast across the tundra zone on the north in the Arctic and then through the taiga zone. Therefore, during quick phase, the three most common vegetation land cover types were forest (44.5±2.9 %, hereafter, percentage of all mean daily positions annotated with the given vegetation type), shrublands (29.9±3 %), and grasslands (24.7±2.8 %, Figure 3a). During slow phase, the three most common vegetation land cover types were grasslands (65.1±6.2 %), croplands (26.9±6.3 %), and forests (4.9±1.4 %, Figure 3b). According to the linear mixed-effects models, the percentage of all vegetation land cover types differed between the slow and quick phases (p<0.001, Table S9), except for the urban lands. Urban lands were more common during the slow than quick phase (Figure 3). However, this type has been annotated for too few birds to make an adequate comparison.

Revised version: During quick phases, Rough-legged buzzards cross the forest zone, while during the slow phase, they migrate within the grassland and cropland zone (Figure 2). Rough-legged buzzards migrated fast across the tundra zone on the north in the Arctic and then through the taiga zone. Therefore, during the quick phase, the three most common vegetation land cover types were forest (44.5±2.9 %, hereafter, percentage of all mean daily positions annotated with the given vegetation type ± sd), shrublands (29.9±3 %), and grasslands (24.7±2.8 %, Figure 2a). During the slow phase, the three most common vegetation land cover types were grasslands (65.1±6.2 %), croplands (26.9±6.3 %), and forests (4.9±1.4 %, Figure 2b). According to the linear mixed-effects models, the percentage of all vegetation land cover types differed between the slow and quick phases (p<0.001, Table S9), except for the urban lands. Urban lands were more common during the slow than quick phase (Figure 2). However, this type has been annotated for too few birds to make an adequate comparison.

Earlier version: Vegetation land cover during quick and slow phases.

Revised version: Vegetation land cover during quick and slow phases of the migration.

Earlier version: a) Quick phase (spring and fall periods together). b) Slow phase (1st and 2nd parts together). c) Migration map.

Revised version: a) Quick phase (spring and fall periods together). b) Slow phase (1st and 2nd parts together). On both a) and b), the bars show the percentage of all mean daily positions annotated with the given vegetation type ± sd. c) Migration map.

Snow cover – the main reason for the dynamic winter range

Earlier version: During the slow phase of the foxtrot migration, Rough-legged buzzards experienced snow cover ranging from 4.8±1.0% in October to 85.2±4.6% in February (Figure 4). If birds spent the winter in the place where they arrived after the fall migration, they would experience snow cover conditions ranging from 4.6±0.6% in October to 99.5±0.1% in February (Figure 4). And if birds fly directly to the southeast and stay there for the whole winter, they would experience snow cover conditions ranging from 1.4±0.2 % in October to 81.1±5.0 % in January (Figure 4). Thus, if birds fly immediately to the southwest and stay there until the end of the winter, they will find conditions with less snow cover in spring (p<0.001, Table S10). And if birds stay where they ended the fall migration, they will find themselves in situations with more snow cover (p<0.001, Table S10). In the latter case, the difference between real and hypothetical situations is not as pronounced (85.2% vs. 99.5%), but more means that snow cover will be close to 100% for several months in this hypothetical situation (Figure 4b).

Revised version: During the slow phase of the migration, Rough-legged buzzards experienced snow cover ranging from 4.8±1.0% in October (hereafter mean±sd) to 85.2±4.6% in February (Figure 3). If birds spent the winter in the place where they arrived after the fall migration, they would experience snow cover conditions ranging from 4.6±0.6% in October to 99.5±0.1% in February (Figure 3b, green line). If birds fly directly to the southeast and stay there for the whole winter, they would experience snow cover conditions ranging from 1.4±0.2 % in October to 81.1±5.0 % in January (Figure 3b, red line). Thus, if birds fly immediately to the southwest and stay there until the end of the winter, they will find conditions with less snow cover in spring (p<0.001, Table S10). And if birds stay where they ended the fall migration, they will find themselves in situations with more snow cover (p<0.001, Table S10). In the latter case, the difference between real and hypothetical situations is not as pronounced (85.2% vs. 99.5%), but more means that snow cover will be close to 100% for several months in this hypothetical situation (Figure 3b).

Earlier version: Snow cover conditions during the slow phase of the foxtrot migration.

Revised version: Snow cover conditions during the slow phase of the migration.

Earlier version: a) 95% Minimum convex polygons (MCPs) of Rough-legged buzzards during winter. Arrows indicate the direction of the movement across months. OCT – October, NOV – November, DEC – December, JAN – January, FEB – February, MAR – March, APR – April. b) Snow cover conditions for the real situation (black) and two hypothetical situations – if birds spend the winter in the place where they arrived after the fall migration (green) and if birds fly directly to the southwest and stay there all winter (red). Dots represent mean values, error lines – standard errors.

Revised version: a) 95% Minimum convex polygons (MCPs) of Rough-legged buzzards during winter. Arrows indicate the direction of the movement across months. OCT – October, NOV – November, DEC – December, JAN – January, FEB – February, MAR – March, APR – April. b) Snow cover conditions for the real situation (black) and two hypothetical situations – if birds spend the winter in the place where they arrived after the fall migration (green) and if birds fly directly to the southwest and stay there all winter (red). Dots represent mean values, error lines – standard deviations.

Removed in the revised version: In some raptor species, adult females disperse further than males (Mearns & Newton 1984; Serrano et al. 2001; Bildstein 2006; Whitfield et al. 2009). While our analysis does not support this (Table S8), it should be acknowledged that our dataset comprises 35 females, eight males, and no juveniles. The unbalanced sex ratio and absence of immature individuals in our sample may have potential influences towards the observed movement patterns in this study.

Discussion

Added in the revised version: Our study identified and characterized a bird migration pattern consisting of an alternation of quick (fall and spring migrations) and slow (non-breeding movements) phases. This migration pattern causes the non-breeding range of birds to shift and become dynamic throughout the season. This has obvious implications for range delineation and assessment of bird population trends.

Added in the revised version: Non-breeding movements and migratory patterns

Added in the revised version: Different bird species employ various wintering strategies, leading to diverse migration patterns. One strategy is residency, where birds arriving at non-breeding grounds remain there throughout the season, undertaking only minor foraging flights (Figure 4a). Their home ranges in non-breeding areas are compatible with those in breeding areas (Kjellén et al. 1997, Alerstam et al. 2006, Newton 2008). Another strategy, initially identified by Moreau (Moreau 1972), is itinerancy, wherein birds fly between several sites separated by distances ranging from ten to several hundred kilometers throughout the non-breeding period, spending 1-2 months at each site (Figure 4b). Most Palearctic-African species adopt this itinerant behavior (Trierweiler et al. 2013, Thorup et al. 2017, Schlaich et al. 2023). The third strategy investigated here involves birds moving slowly and steadily throughout the non-breeding season, initially in one direction and then in the opposite direction. In 2009, Strandberg et al. used GPS transmitters to track Common buzzards (Buteo buteo). Despite some transmitters in their study ceasing to function mid-winter, the authors observed a phenomenon they termed “prolonged autumn migration,” wherein birds arriving at their wintering grounds did not halt migration but continued moving at a slower pace (Strandberg et al. 2009). In our study, we observed this phenomenon throughout the annual cycle. We demonstrated that, unlike itinerancy, movements in the non-breeding area are essentially a continuation of migration at a slower pace, influenced by external factors such as snow cover dynamics. Considering this behavior as an extension of migration and examining the complete life cycle of such species, it involves a quick phase during the transition between breeding and non-breeding areas, followed by a slow phase of directed and seasonal movement within the non-breeding area, and finally, another quick phase towards the breeding area. Drawing an analogy to the alternating quick and slow movements of the foxtrot dance, we propose the term ‘foxtrot migration’ for this seasonal movement to offer a concise and easily understandable description (Figure 4c). Consequently, we suggest referring to the non-breeding range of species exhibiting itinerancy or foxtrot migration as the ‘dynamic range’ (Figure 4b, c).

Added in the revised version: Winter strategies and non-breeding range scheme.

Added in the revised version: a) Residency, b) Itinerancy, c) Foxtrot migration. The color bar along the y-axis corresponds to the color coding of the habitats of the chart on the right.

Added in the revised version: Debates may arise regarding introducing a new term for this phenomenon and whether it should fall under the general term “nomadism” or be categorized as a modified form of another migratory pattern, such as “slow directional itinerancy.” We assert that behavioral pattern names should aid in understanding movement strategies and avoid confusion. One might ask whether it is necessary to distinguish between migration and wintering at all, given that migration is essentially a matter of tracking spatial and temporal changes in favorable conditions. Authors studying Montagu’s harriers (Circus pygargus) argue that such generalization hampers understanding of wintering strategies (Schlaich et al. 2023). Instead, Schleich et al. (2023) propose explicitly considering the functions of animal-used sites, particularly their contributions to “wintering” and “refueling for migration.” For instance, cuckoos (Canorus sp.) utilize stopovers for both peak foraging conditions and refueling for long flights to subsequent sites (Thorup et al. 2017). Conversely, Montagu’s harriers exhibit negligible refueling between sites due to short distances and energy-efficient soaring flight (Trierweiler et al. 2013, Schlaich et al. 2023). In our case, Rough-legged buzzards do not make several long stopovers, unlike species experiencing itinerancy, but move gradually in one direction and then return during the whole season (Figure 1a). As Schleich et al. (2023) suggested, examining the functions of animal-used sites reveals another difference between these strategies. The difference is that itinerant birds wait for better conditions at subsequent sites. In contrast, in our case, the southeastern areas that birds reach after several months are suitable and offer better conditions than the northeastern areas, as they have less snow throughout the whole non-breeding season (Figure 3). All these differences highlight the discrepancy between the wintering strategies referred to as itinerancy and those described in this study. Therefore, using another term - foxtrot migration - for this phenomenon is more appropriate.

Added in the revised version: The dynamic of the foxtrot migration

Removed in the revised version: Our study confirmed the existence of a previously hypothesized bird migration pattern (García & Arroyo 1998) characterized by alternating quick and slow phases (Figure 1). Following the breeding season, birds quickly traverse unfavorable habitats and proceed slowly through winter in favorable environments influenced by external factors. Subsequently, they quickly cross unfavorable terrain once more to reach the breeding region in spring (Figure 2). This pattern, termed ‘foxtrot migration’, distinguishes itself from ‘itinerancy’ by featuring directional and continuous movements. Additionally, it differs from ‘forage-on-fly migration’ as it occurs during the non-breeding period. We propose referring to the non-breeding area as ‘dynamic area’, reflecting its temporal heterogeneity, where the region with the highest bird density continually changes throughout the non-breeding period.

Earlier version: Our study affirmed the presence of foxtrot migration and a dynamic non-breeding range in Rough-legged Buzzards. The taiga zone posed as unfavorable habitat during the quick phase of foxtrot migration (Figure 3), given the difficulty for Rough-legged Buzzards to locate open areas for hunting. Conversely, the grassland and cropland zone served as favorable habitats during the slow phase throughout the entire non-breeding period (Figure 3), offering numerous open areas for hunting. Snow cover was the external factor driving their continual 1000 km southwest movement during winter (Figure 4). Our analysis revealed that if Rough-legged Buzzards remained at their fall migration endpoint without moving southwest, they would encounter 14.4% more snow cover (99.5% vs. 85.1%, Figure 4). Although this difference may seem small (14.4%), it holds significance for rodent-hunting birds, distinguishing between complete and patchy snow cover. Simultaneously, if Rough-legged Buzzards immediately flew to the southwest and stayed there throughout winter, they would experience 25.7% less snow cover (57.3% vs. 31.6%, Figure 4). Despite a greater difference than in the first case, it doesn’t compel them to adopt this strategy, as it represents the difference between various degrees of landscape openness from snow cover. This indicates ecological tipping points in the environment that are not predicted by a slight linear increase in effect size.

Revised version: Our study affirmed the presence of foxtrot migration and a dynamic non-breeding range in Rough-legged buzzards. The taiga zone posed as an unfavorable habitat during the quick phase of foxtrot migration (Figure 2), given the difficulty of Rough-legged buzzards locating open areas for hunting in this habitat. Conversely, the grassland and cropland zone served as favorable habitats during the slow phase throughout the entire non-breeding period (Figure 2), offering numerous open areas for hunting. Snow cover was the external factor driving their continual 1000 km southwest movement during winter (Figure 3). Our analysis revealed that if Rough-legged buzzards remained at their fall migration endpoint without moving southwest, they would encounter 14.4% more snow cover (99.5% vs. 85.1%, Figure 3). Although this difference may seem small (14.4%), it holds significance for rodent-hunting birds, distinguishing between complete and patchy snow cover. Simultaneously, if Rough-legged buzzards immediately flew to the southwest and stayed there throughout winter, they would experience 25.7% less snow cover (57.3% vs. 31.6%, Figure 3b). Despite a greater difference than in the first case, it doesn’t compel them to adopt this strategy, as it represents the difference between various degrees of landscape openness from snow cover. However, it is noteworthy that one bird followed this strategy by flying directly to the southwest and remaining there throughout the winter (Figure S1), exhibiting a residency strategy. This observation suggests that the distinction between winter strategies is flexible.

Added in the revised version: African-Palearctic migrants do not strictly adhere to a binary division between residency and itinerancy but demonstrate a continuum between these two strategies. Many species in Africa exhibit a mixture of both strategies in varying proportions; see Schlaich et al. (2023) and reference therein. Similarly, in the mid-latitudes, species may not only show any of the three strategies but also a considerable degree of individual and regional variation. For example, Rough-legged buzzards from western Arctic regions that overwinter in mid-latitude areas where snow cover dynamics are less pronounced than in eastern regions are likely to have a higher proportion of birds with a resident strategy. At the same time, Rough-legged buzzards in North America also exhibit a foxtrot migration pattern, although with variations. There, birds make a direct and rapid migration, usually across the boreal forest, followed by slower, more facultative movements once birds get south of the boreal. However, there is a lot of variation among individuals and even regions; there are more non-breeding slow movements in eastern North America but less in western North America (Neil Paprocki pers. comm.). Thus, Rough-legged buzzards in North America lie along a gradient between foxtrot migration and residency, likely influenced by different snow cover dynamics on the east and west coasts.

Added in the revised version: The foxtrot migration pattern is expected to be observed in many migratory species that exhibit distinct seasonal cycles in their non-breeding range. This migration pattern is expected to be prevalent for species living in regions where snow cover is a significant determinant of food availability, which is the case in large parts of the mid-latitudes. Therefore, understanding bird migration patterns is critical to accurately mapping ranges and assessing population trends in this region.

Added in the revised version: Mapping ranges and assessing population trends

Removed in the revised version: This pattern is likely to be observed in many migratory species with distinct seasonal cycles in their non-breeding range. In our study, we focused on an Arctic species that resides in the mid-latitudes during the non-breeding period. In general, foxtrot migration is likely to be a common feature among species where snow cover significantly influences food availability, especially those specializing in rodents. Originally proposed for birds in Africa, where seasonal changes regulate food availability, it’s likely that many insectivorous birds that spend the non-breeding season in Africa also exhibit this form of migration. This type of migration and dynamic non-breeding range is widespread both geographically and taxonomically, with obvious conservation implications.

Removed in the revised version: The conservation implications

Earlier version: The conservation implications of our study are twofold: 1) the use of mid-winter (Christmas) bird surveys to determine non-breeding range may yield inaccurate results for species with foxtrot migrations (Figure 5a), and 2) declines in abundance within a particular segment of the non-breeding range may indicate changes in range dynamics rather than widespread declines in species abundance (Figure. 5b).

Revised version: The implementation of our study is twofold: 1) the use of mid-winter bird surveys to determine non-breeding range may yield inaccurate results for species with dynamic range (Figure 5a), and declines in abundance within a particular segment of the non-breeding range may indicate changes in range dynamics rather than widespread declines in species abundance (Figure. 5b).

Added in the revised version: Mapping ranges and assessing population trends for the species with the dynamic non-breeding ranges.

Removed in the revised version: The use of foxtrot migration research for the assessment of the non-breeding range and conservation status of a particular species.

Earlier version: a) Christmas counts, often used to determine winter range, may give a misleading representation of non-breeding range for species with foxtrot migration. b) Changing the conservation status of a species based on reduced abundance in a particular area may not be appropriate for species with foxtrot migrations. Red – breeding range, blue – winter range. Latin numbers indicate the month the birds are present in the area.

Revised version: a) Mid-winter counts, often used to determine winter range, may give a misleading representation of the non-breeding range for species with dynamic range. b) Changing the conservation status of a species based on reduced abundance in a particular area may not be appropriate for species with dynamic range. Red – breeding range, blue – winter range.

Earlier version: Christmas Bird Counts provide insight into the mid-winter distribution of a species. For example, based on our study of the Rough-legged Buzzard, the species is predominantly present in the southwestern portion of its non-breeding range during this period, with only a small proportion present in the rest of the region. As a result, a map of the non-breeding range may inaccurately depict the species as common in the southwest and extremely rare in the northeast. This map would be inaccurate because, during the entire non-breeding period, Rough-legged Buzzards spend both fall and spring in the northeastern part and only mid-winter in the southwesternmost part.

Revised version: In North America and Europe, the number of non-breeding birds is typically estimated once a season in mid-winter. In North America, these estimates are made during the Christmas counts, usually once in late December (National Audubon Society 2024), while in Europe, they are made during the IWC counts, usually once in January or February (Wetlands International 2024). This approach can lead to errors in range estimation for birds with dynamic ranges. For such species, habitats occupied by birds in fall and spring will be listed as “uncommon” at best, while habitats occupied by birds only in mid-winter will be listed as “common.” However, the situation is the opposite if we consider the time birds spend in these habitats (Figure 5a). For example, based on our study of the Rough-legged buzzard, during mid-winter, the species is predominantly present in the southwestern portion of its non-breeding range, with only a tiny proportion present in the rest of the region. As a result, a map of the non-breeding range may show the species as “common” in the southwest and “uncommon” in the northeast. This map would be inaccurate because, during the entire non-breeding period, Rough-legged buzzards spend both fall and spring in the northeastern part and only mid-winter in the southwestern part.

Earlier version: To address this, continuous year-round GPS tracking of the species provides a means to track bird locations throughout the non-breeding season, facilitating the creation of accurate distribution maps, particularly for species that exhibit foxtrot migrations and dynamic ranges. We advocate representing temporal heterogeneity (range dynamics) on maps as distinct zones, denoting periods when the species is abundant in a given area. To distinguish temporal from spatial heterogeneity, we recommend using lines to delineate the boundaries of these zones, rather than color shading, and incorporating numbers to denote the months of species abundance in a given zone (Figure 5a). We suggest ecologists include dynamic non-breeding ranges in descriptions and range maps for foxtrot migratory bird species.

Revised version: To address this, continuous year-round GPS tracking of the species provides a means to track bird locations throughout the non-breeding season, facilitating the creation of accurate distribution maps, particularly for species that exhibit dynamic ranges. We advocate representing temporal heterogeneity (range dynamics) on maps as distinct zones, denoting periods when the species is abundant in a given area. To distinguish temporal from spatial heterogeneity, we recommend using lines to delineate the boundaries of these zones, rather than color shading, and incorporating numbers to denote the months of species abundance in a given zone (Figure 5a). We suggest ecologists include dynamic non-breeding ranges in descriptions and range maps for foxtrot migratory and itinerant bird species.

Earlier version: Population counts for a species are often limited to a portion of its range. Therefore, conclusions about conservation status drawn from such counts may be misleading. A decline in abundance within a particular portion of the non-breeding range may indicate changes in range dynamics rather than a general decline in the species. For example, climate change may affect snow cover dynamics, reducing its intensity in northern regions. As a result, species whose range dynamics depend on snow cover may choose to remain in the northern areas and not migrate as far south as they traditionally have. Despite this shift in range dynamics, overall species abundance may remain unchanged. Similar patterns have affected Rough-legged Buzzards in some areas of the European non-breeding range.

Revised version: Population counts for a species are often limited to a portion of its range. Therefore, conclusions about conservation status drawn from such counts may be misleading. A decline in abundance within a particular portion of the non-breeding range may indicate changes in range dynamics rather than a general decline in the species (Figure 5b). For example, climate change may affect snow cover dynamics, reducing its intensity in northern regions. As a result, species whose range dynamics depend on snow cover may choose to remain in the northern areas and not migrate as far south as they traditionally have. Despite this shift in range dynamics, overall species abundance may remain unchanged. Similar patterns have affected Rough-legged buzzards in some areas of the European non-breeding range.

Earlier version: A 2022 Dutch study found a decline in wintering Rough-legged buzzards over the last 40 years (Hornman, Boele & van Winden 2022). On the one hand, this may represent a conservation concern. On the other hand, applying the rationale of the foxtrot migration, the apparent local decline may simply be attributed to climate change, resulting in less comprehensive snow coverage in the northeastern wintering areas of Rough-legged buzzards relative to the Netherlands. Such a shift in snow coverage makes it less probable for the birds to migrate to the Netherlands for overwintering. This proposition is further supported by a study of the winter population dynamics of Rough-legged buzzards in the Netherlands in 2011, showing that the main winter population peak occurred in late December, with many birds migrating (Vansteelant, Faveyts & Buckens 2011). The authors also found that the main migration occurred after heavy snowfall in northern Europe, supporting our foxtrot migration explanation for this decline. Therefore, investigating the dynamic range and foxtrot migration is critical to understanding a species’ range and effectively assessing its conservation status.

Revised version: A 2022 Dutch study found a decline in wintering Rough-legged buzzards over the last 40 years (Hornman et al. 2022). On the one hand, this may represent a conservation concern. On the other hand, applying the rationale of the dynamic range, the apparent local decline may be attributed to climate change, resulting in less comprehensive snow coverage in the northeastern wintering areas of Rough-legged buzzards relative to the Netherlands. Such a shift in snow coverage makes it less probable for the birds to migrate to the Netherlands for overwintering. This proposition is further supported by a study of the winter population dynamics of Rough-legged buzzards in the Netherlands in 2011, showing that the main winter population peak occurred in late December, with many birds migrating (Vansteelant et al. 2011). Vansteelant et al. (2011) also found that the main migration occurred after heavy snowfall in northern Europe, supporting our foxtrot migration explanation for this decline. Therefore, investigating the dynamic range is critical to understanding a species’ range and effectively assessing its conservation status.

Added in the revised version: Conclusion

Added in the revised version: Our study sheds light on the overwintering strategy of the Rough-legged buzzard, an Arctic breeder that migrates to mid-latitudes during the non-breeding season. We have observed a migration pattern in which these birds travel about 1000 km round trip in response to changing snowline conditions, which we call the ‘foxtrot migration’ for its similarity to the alternating quick and slow movements of the foxtrot dance. Our results highlight the importance of this migration pattern and the concept of ‘dynamic range’ in revealing potential inaccuracies in range mapping from single midwinter surveys and suggest that species distribution maps should represent temporal variations in range dynamics as separate zones. Understanding and accounting for this migratory behavior is critical for accurately assessing bird populations in non-breeding areas. This is especially important for species such as Arctic nesting birds, where population assessments are primarily conducted in non-breeding ranges due to the logistical and political challenges of the Arctic.

Removed in the revised version: Significant environmental changes, including changes in snowpack dynamics, have occurred in recent decades that may have significant impacts on bird migration (Tucker et al. 2018; Sumasgutner et al. 2021). In order to understand the effects of global environmental change on bird migration behavior, there is an urgent need for in-depth investigations into the intricate relationship between migration patterns and evolving environmental factors. The advent of new technologies in recent decades has expanded our understanding of bird migration, enabling the identification of novel behavioral patterns and facilitating the study of their susceptibility to various environmental influences. This knowledge is becoming increasingly urgent in the face of rapid climate change and anthropogenic habitat alteration, which pose threats to both wildlife and human populations.

Removed in the revised version: Conclusions

Removed in the revised version: Our study has identified and characterized a new pattern of migratory behavior, the ‘foxtrot migration’, along with the associated concept of ‘dynamic range’. This discovery has significant implications for conservation strategies and adequate representation of non-breeding habitats. As animal tracking technology advances, we expect to discover more new details about animal migrations and, consequently, new patterns of migratory behavior. With these technological developments, we are gaining a clearer understanding of the complex patterns and behaviors that drive animal migrations, providing valuable insights for conservation efforts and deepening our understanding of the natural world.

Supporting information